---
title: 'Introducción a los fundamentos de R'
description: 'En este capítulo, dará sus primeros pasos en Python, aprenderás cómo usar la consola como calculadora y cómo asignar variables. También aprenderá acerca de los principales tipos de datos utilizados en Python. Vamos a empezar!'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
